<problem display_name="Sample Problem: Text Logger" markdown="null" max_attempts="5" rerandomize="onreset" show_reset_button="true" showanswer="finished">
  <script type="loncapa/python">

# The CDATA declaration lets us use less-than signs
# without edX thinking that we're starting an XML tag.

from python_lib import HXGraders

def answercheck(e, ans):

    options = {
        'min_length': 10   # This is the minimum length in characters of the response.
    }

    # This pulls from the python_lib.zip file in Files &amp; Uploads.
    return HXGraders.textResponseGrader(ans, options)


</script>
  <script type="text/javascript">
$(document).ready(function(){
	console.log('outer ready');
});

function logThatThing(ThatThing){
	console.log(JSON.stringify(ThatThing));
  // Call the edX logging code.
	Logger.log("harvardx.StudioAdv.text_logger", ThatThing);
}
</script>
  <p>Let's walk through a really simple problem: this one! It lets you type things in below and submit them. There's no grading; you get 100% as long as you type at least 10 characters. We use this for journaling, or just to log a student's response for later analysis.</p>
  <p>Play around if you want, then go below and get the files.</p>
  <customresponse cfn="answercheck">
    <jsinput get_statefn="textlog.getState" gradefn="textlog.getGrade" height="120" html_file="/static/textlog.html" set_statefn="textlog.setState" sop="false" title="Text Response Exercise" width="800"/>
  </customresponse>
  <demandhint>
    <hint>Just a reminder, hints work in these problems too if you want!</hint>
    <hint>If you need the <a href="https://edx.readthedocs.io/projects/edx-partner-course-staff/en/latest/exercises_tools/text_input.html#configure-hints-in-the-advanced-editor" target="_blank">XML formatting for on-demand hints</a>, it's the same as in any other problem type. Follow dat link.</hint>
  </demandhint>
  <solution>
    <div class="detailed-solution">
      <p>Explanation</p>
      <p>You can put some useful commentary into the answer box, or you can just set the problem to never show the answer. Do whatever makes sense for your situation. </p>
    </div>
  </solution>
</problem>
